<a href="https://colab.research.google.com/github/gt-cse-6040/skills_oh_week_05/blob/main/week02_session02_NB02_dict_diffs_topValue_edited.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

_Main topics covered during today's session:_

Prior NB:

1. **Counting values in a dictionary**
    
    
This NB:

2. **Comparing dictionaries, and finding the highest/lowest value.**


Following NBs:

3. **Inverting a dictionary**

4. **Teaser -- Floating point values**

# Comparing dictionaries and finding the highest/lowest values.

### This notebook was part of MT1 in the first iteration of CSE6040, when the MT's were a number of smaller and independent notebooks.

#### We will be doing Exercises 1 and 2, to demonstrate comparing two dictionaries and finding the highest/lowest value in a dictionary.

# Problem 7

**Letter frequencies.** This problem has three (3) exercises worth a total of ten (10) points.

Letter frequency in text has been studied in cryptoanalysis, in particular frequency analysis. Linguists use letter frequency analysis as a rudimentary technique for language identification, where it's particularly effective as an indicator of whether an unknown writing system is alphabetic, syllablic, or ideographic.

Primarily, three different ways exist for letter frequency analysis. Each way generally results in very different charts for common letters. Based on the provided text, the first method is to count letter frequency in root words of a dictionary. The second way is to include all word variants when counting, such as gone, going and goes and not just the root word go. Such a system results in letters like "s" appearing much more frequently. The last variant is to count letters based on their frequency in the actual text that is being studied. 

For more details, refer to the link: 
https://en.wikipedia.org/wiki/Letter_frequency

In this problem, we will focus on the 3rd methodology.

**Exercise 0** (2 points). First, given a string input, define a function  `preprocess` that returns a string with non-alphabetic characters removed and all the alphabets converted into a lower case. 

For example, 'We are coding letter Frequency! Yay!" would be transformed into "wearecodingletterfrequencyyay"

```python
def preprocess(S):
    ###
    ### YOUR CODE HERE
    ###

    ret_string = ''
    for i in S:
        if i.isalpha():
            ret_string = ret_string + i
    
    ret_string = ret_string.lower()
    return ret_string

```

<!-- grade_id: valid_string, points: 2, grade: true, solution: false -->
```python
# Test cell: valid_string
import random, string

N_str = 100 #Length of random string

def generate_str(n):
    random_str = ''.join(random.choice(string.ascii_lowercase + string.ascii_uppercase + string.digits + string.punctuation) for _ in range(n))
    return random_str

def check_preprocess_str(n):
    random_str = generate_str(n)
    print("Input String: ",random_str)
    assert preprocess(random_str).islower() == True
    assert preprocess(random_str).isalpha() == True
    print("|----Your function seems to work correct for the string----|"+"\n")

check_preprocess_str(N_str)
check_preprocess_str(N_str)
check_preprocess_str(N_str)

print("\n(Passed)!")
```

Output:
```
Input String:  tn,J^M3"?b9y*Vkxv>{J^1>Z+M<W\IvQ?:Qt^|5;(#gW3,{*rZm~4%{h+ZtMvtNbJVPpK:!RS>R@r}1D"e6gq,ZF"BW5LUr|aaxG
|----Your function seems to work correct for the string----|

Input String:  9rGt@J[l,)KquDt^Z?2Ju`V?Adz_!69A/<;UpGH5U)m6*{5Ki\x>WSh97Sb&kM^fNW,;(jG=8px_EC(]7^1.?nO)P$VAu|9v>+JQ
|----Your function seems to work correct for the string----|

Input String:  !Z]G]G!:K7ieH["sN*^,[;qs4^k%@B!0v^{jkL&Cfgl%3^$o-vKi0~5rzagfXWVJhk|U`F3LIUszZl.6lOZh8Mf|kJs#GK\qQ<Ik
|----Your function seems to work correct for the string----|


(Passed)!
```

**Exercise 1** (4 points). With the necessary pre-processing complete, the next step is to write a function `count_letters(S)` to count the number of occurrences of each letter in the alphabet.  

You can assume that only letters will be present in the input string. It should output a dictionary and if any alphabet (a-z) is missing in the input string, it should still be a part of the output dictionary and its corresponding value should be equal to zero.


```python
def count_letters(S):
    ###
    ### YOUR CODE HERE
    ###

    # Larry's code, not the sample solution
    import string
    # create dictionary of letters with 0 values
    alphabet = dict.fromkeys(string.ascii_lowercase, 0)
#     print(alphabet)
    
    # get the letters and their count from the input word S
    from collections import Counter
    count = Counter(S)
#     print(count)
    
    # compare the two dictionaries and add the letters that
    # are not in the count dictionary
    for key,value in alphabet.items(): 
        if not key in count:
            count.update({key:value})
            
    return count      
            
# https://stackoverflow.com/questions/453576/is-there-a-fast-way-to-generate-a-dict-of-the-alphabet-in-python
# https://www.geeksforgeeks.org/python-difference-in-keys-of-two-dictionaries/


#count_letters('abcdef')
```

<!-- grade_id: count_letters, points: 4, grade: true, solution: false -->
```python
# Test cell: count_letters
import collections

N_processed_str = 100

def generate_processed_str(n):
    random_processed_str = ''.join(random.choice(string.ascii_lowercase) for _ in range(n))
    return random_processed_str

def check_count_letters(S):
    print("Input String: ",S)
    random_char = chr(random.randint(97,122))
    print("Character frequency evaluated for: ", random_char)
    if(random_char in S):
        assert count_letters(S)[random_char] == collections.Counter(S)[random_char]
        print("|----Your function seems to return correct freq for the char----|"+"\n")
    else:
        assert count_letters(S)[random_char] == 0
        print("|----Your function seems to return correct freq for the char----|"+"\n")
        
check_count_letters(generate_processed_str(N_processed_str))
check_count_letters(generate_processed_str(N_processed_str))
check_count_letters(generate_processed_str(N_processed_str))
print("\n(Passed)!")
```

Output:
```
Input String:  xenotfvkyapzpnprwefejolnuqvpzgaxamasorfpagnsnoriqkwqcdqdhxbvebnlzqbeldskmneatorsvyplegxptbacmkknsbou
Character frequency evaluated for:  x
|----Your function seems to return correct freq for the char----|

Input String:  owoalnbxkuqlpwmakocgktwbzgavtbgorvghtmfsyzepvijpkzzvbpezrvhoytiobdsbrazgzihtdrnbhmqngmehmfuvdnamknpi
Character frequency evaluated for:  w
|----Your function seems to return correct freq for the char----|

Input String:  zjzwstninndhcrdyhxebhswdbvqgjfmqndeancdksdjpuvpjvtlekrtjorfjjnqabiiljordvziajfpfoynvodrnijqrtwkjztqy
Character frequency evaluated for:  h
|----Your function seems to return correct freq for the char----|


(Passed)!
```

**Exercise 2** (4 points). The next step is to sort the distribution of a dictionary containing all the letters in the alphabet as keys and number of occurrences in text as associated value. 

Sorting should be first done in decreasing order by occurrence count and for two elements with same count, the order should be alphabetic. The function  `find_top_letter(d)` should return the 1st character in the order.

```python
def find_top_letter(d):
    ###
    ### YOUR CODE HERE
    ###

    # Larry's code, not the sample solution can be viewed in the Colab link to this NB 
    sorted_d = sorted(d.items(), key=lambda x:(-x[1], x[0]))
    # sorted_d = sorted(sorted_d, key=lambda x: x[1], reverse=True)#list of tuples sorted as in the prompt
    
    return sorted_d[0][0]
    
'''
https://stackoverflow.com/questions/15371691/how-to-sort-a-dictionary-by-value-desc-then-by-key-asc


'''
find_top_letter(count_letters('accccdeffffbbbb'))
```

Output:
```
'b'
```

<!-- grade_id: highest_freq_letter, points: 4, grade: true, solution: false -->
```python
# Test cell: highest_freq_letter

def create_random_dict():
    max_char_value = random.randint(5, 20)
    random_dict = {c:random.randint(0,max_char_value-1) for c in string.ascii_lowercase}
    random_letter1, random_letter2 = random.sample(string.ascii_lowercase, 2)
    random_dict[random_letter1], random_dict[random_letter2] = max_char_value, max_char_value
    if(random_letter1 < random_letter2):
        return random_letter1, random_dict
    else:
        return random_letter2, random_dict

def check_top_letter():
    top_letter, random_dict = create_random_dict()
    user_letter = find_top_letter(random_dict)
    assert user_letter == top_letter
    print("Input Dictionary: ", random_dict)
    print("Your function correctly returned most frequent letter: {} \n".format(user_letter))
    
check_top_letter()
check_top_letter()
check_top_letter()
print("\n(Passed)!")
```

Output:
```
Input Dictionary:  {'a': 3, 'b': 7, 'c': 14, 'd': 11, 'e': 10, 'f': 7, 'g': 7, 'h': 15, 'i': 7, 'j': 12, 'k': 10, 'l': 16, 'm': 13, 'n': 0, 'o': 1, 'p': 14, 'q': 3, 'r': 14, 's': 16, 't': 4, 'u': 15, 'v': 10, 'w': 15, 'x': 4, 'y': 14, 'z': 5}
Your function correctly returned most frequent letter: l 

Input Dictionary:  {'a': 4, 'b': 4, 'c': 3, 'd': 6, 'e': 9, 'f': 12, 'g': 0, 'h': 0, 'i': 17, 'j': 14, 'k': 7, 'l': 0, 'm': 7, 'n': 17, 'o': 15, 'p': 20, 'q': 19, 'r': 14, 's': 6, 't': 6, 'u': 20, 'v': 1, 'w': 3, 'x': 3, 'y': 0, 'z': 14}
Your function correctly returned most frequent letter: p 

Input Dictionary:  {'a': 13, 'b': 3, 'c': 5, 'd': 16, 'e': 5, 'f': 11, 'g': 9, 'h': 14, 'i': 14, 'j': 5, 'k': 2, 'l': 9, 'm': 15, 'n': 1, 'o': 12, 'p': 4, 'q': 4, 'r': 2, 's': 2, 't': 11, 'u': 6, 'v': 16, 'w': 7, 'x': 8, 'y': 11, 'z': 12}
Your function correctly returned most frequent letter: d 


(Passed)!
```

**Fin!** You've reached the end of this problem. Don't forget to restart the kernel and run the entire notebook from top-to-bottom to make sure you did everything correctly. If that is working, try submitting this problem. (Recall that you *must* submit and pass the autograder to get credit for your work!)